```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment3.ipynb")
```

# Assignment 3 - Data Breaches

The questions in each assignment will be mostly autograded, but you will **not** see all the autograded tests in the notebook. After the deadline has passed, many additional tests will run on the file you submitted on Gradescope and you will be able to view your score online. In addition, please note that while we do NOT explicitly tell you to add chart titles and axis labels, we strongly encourage you to do so for the visualizations you create. In addition to the autograding, the TAs, will manually grade your use of chart titles and appropriate axis labels on this assignment.

## Assessment
7% of your Overall course Grade

## Due Date
February 21st (end of day 11:59pm)

## Submission instructions

<div class="alert alert-info" style="color:black">
<ul>
  <li>Download this lab from the Jupyter Book web page by clicking the download symbol at the top right of the page and selecting the <code>.ipynb</code> format.</li>
  <li>To submit this lab, answer all the questions and then upload the completed lab to Gradescope.</li>
    <ul>
      <li>Before submitting, make sure you restart the kernel and rerun all cells (click the ▶▶ button).</li>
    </ul>
  <li>Don't change any variable names that are given to you, don't move cells around, and don't include any code to install packages in the notebook.</li>
  </li>
</ul>
</div>

```python
# Run this cell to ensure that altair plots show up on gradescope
# We will talk more about what these lines do later in the course
import altair as alt

# Handle large data sets without embedding them in the notebook
alt.data_transformers.enable('data_server')
# Include an image for each plot since Gradescope only supports displaying plots as images
alt.renderers.enable('mimetype')
```

### Question 1

<div class="alert alert-info" style="color:black">

Which sectors have had the most data breaches?
To find out, complete the following tasks:

1. Run the cell that reads in the data and study the dataframe to get familiar with the column values.
2. Create a bar chart to determine which sectors have the most data breaches (i.e., this is the count of data breaches not the number of records lost)
3. Choose a suitable orientation for what goes on the x and y axes.
4. Sort the bars based on their value, with the largest bar the closest to the x-axis line.
    
</div>

_Points:_ 6

```python
# Provided code to read in the data
import pandas as pd

url = 'https://raw.githubusercontent.com/kemiolamudzengi/dsci-320-datasets/main/breaches.csv'
breaches = pd.read_csv(url, parse_dates=['date'])
breaches
```

```python
bar_sector_count = ...

# Show the plot
bar_sector_count
```

```python
grader.check("q1")
```

### Question 2

<div class="alert alert-info" style="color:black">

Out of all the data breaches from each sector, how many are there of each data sensitivity level?
To find out, complete the following task:
    
Building off of the your answer for Q1, add information about how sensitive the data breaches in each category were,
   by creating a stacked bar chart
   where each section of the bar is encoded with color according to the data sensitivity.
   Make sure that the data sensitivity is encoded as an ordinal color variable,
   to reflect that there is an order between the five levels of sensitivity
   (1 = least sensitive, 5 = most sensitive).

</div>

_Points:_ 7

```python
bar_sector_sensitivity_count = ...

# Show the plot
bar_sector_sensitivity_count
```

```python
grader.check("q2")
```

### Question 3

<div class="alert alert-info" style="color:black">

It can often be difficult to compare the length of segments in a stacked bar chart,
since they do not start at the same baseline.
Another way of visualizing the information in the previous plot,
is to create a scatterplot with counts of the combinations of each sector and data sensitivity category:

1. Create a scatterplot
   with the sectors on the y-axis
   and the data sensitivity level on the x-axis
   (encoded as an ordinal variable).
2. Encode the counts of each combination as the size of the points.
3. Sort the y-axis based on the counts (the size encoding).
4. Add a tooltip with the exact count of each categorical combination.

</div>

_Points:_ 8

```python
point_sector_sensitivity_count = ...

# Show the plot
point_sector_sensitivity_count
```

```python
grader.check("q3")
```

### Question 4

<div class="alert alert-info" style="color:black">

Another way of visualizing the information in the previous plot is to create a heatmap with counts of the combinations of each sector and data sensitivity category:

1. Create a heatmap
   with the sectors on the y-axis
   and the data sensitivity level on the x-axis
   (encoded as an ordinal variable).
2. Encode the counts of each combination as the color of the rectangle
3. Sort the y-axis based on the counts.
4. Add a tooltip with has the sector, data sensitity and count

</div>

_Points:_ 14

```python
heatmap_sector_sensitivity_count = ...

# Show the plot
heatmap_sector_sensitivity_count
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

### Question 5

<div class="alert alert-info" style="color:black">

Compare the last three visualizations that you created. Which is most effective at answering the questions below.
Answer the questions and use the theoretical concepts you have been exposed (in the course) to explain why that visualization is more effective than the others

1. Which sector has the highest number of breaches in our dataset?
2. What proportion of leaked government data is most sensitive (i.e. level 5 sensitivity)?
</div>

_Points:_ 0

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 6

<div class="alert alert-info" style="color:black">

Next, let's explore how the size of the breaches have changed over time.
To make the plot easier to interpret,
we will focus on the four sectors with the most data breaches.
The necessary data wrangling for this has already been done for you,
and you need to use the wrangled data to complete the following tasks:

1. Create a scatter plot with the date on the x-axis
   and the records lost on the y-axis.
2. Let the size of the points indicate the sensitivity of the breached data.
3. Let the color indicate which sector the point belongs to.
4. Add a tooltip that displays the source of the breach
  when the point is hovered over with the mouse pointer.

</div>

_Points:_ 10

```python
# Provided data wrangling
top_sectors = breaches['sector'].value_counts().nlargest(4).index.tolist()
top_breaches = breaches.query('sector in @top_sectors')
```

```python
scatter_breaches_over_time = ...

# Show the plot
scatter_breaches_over_time
```

```python
grader.check("q6")
```

### Question 7

<div class="alert alert-info" style="color:black">

For the next 3 questions, let's return to the original dataset (i.e., breaches)
Next, let's use a sorted bar chart to determine which method resulted in the most records lost

1. Create a bar chart with the sum of records lost encoded on the x-axis and the method encoded on the y-axis
2. Sort such that the method with the highest number of records lost is furthest away from the x-axis
</div>

_Points:_ 9

```python
bar_method_records_lost = ...

# Show the plot
bar_method_records_lost
```

```python
grader.check("q7")
```

### Question 8

<div class="alert alert-info" style="color:black">

Next, let's use a **normalized** stacked bar to get a sense of the sensitivity distribution for each method.
Create a stacked bar chart with
1. The sum of records lost encoded on the x-axis
2. The method on the y-axis
3. Sensitivity on the color channel
</div>

_Points:_ 12

```python
sbar_method_records_sensitivity = ...

# Show the plot
sbar_method_records_sensitivity
```

```python
grader.check("q8")
```

### Question 9

<div class="alert alert-info" style="color:black">

You have done alot, let's create one more visualization.
Create a scatterplot with using the dataset for the given sectors
1. sector encoded on y channel
2. method encoded on x channel
3. sum of records_lost with size
4. data sensitivity with color
</div>

_Points:_ 8

```python
# Provided Wrangling
sector_list = ['app', 'finance', 'government', 'web', 'military']
selected_sectors = breaches.loc[(breaches['sector'].isin(sector_list))]

```

```python
scatter_some_sectors = ...

# Show the plot
scatter_some_sectors
```

```python
grader.check("q9")
```

<!-- BEGIN QUESTION -->

### Question 10

<div class="alert alert-info" style="color:black">

For each of the visualizations you have created (i.e., 8 in total), using the task abstraction framework presented in class, write out one task that each viz is effective at.
 In addition write out one thing your learned about data breaches from this dataset/assignment </div>

_Points:_ 0

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<div class="alert alert-danger" style="color:black">
    
**Restart and run all before submitting**
    
Before submitting your notebook.
don't forget to run all cells in your notebook
to make sure there are no errors
and so that the TAs can see your submission on Gradescope.
You can do this by clicking the ▶▶ button
or going to `Kernel -> Restart Kernel and Run All Cells...` in the menu.
This is not only important for DSCI 320,
but a good habit you should get into before ever committing a notebook to GitHub,
so that your collaborators can run it from top to bottom
without issues.
</div>

Tests for q1 (points: 6)
```
>>> assert bar_sector_count is not None
>>> assert isinstance(bar_sector_count, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q2 (points: 7)
```
>>> assert bar_sector_sensitivity_count is not None
>>> assert isinstance(bar_sector_sensitivity_count, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q3 (points: 8)
```
>>> assert point_sector_sensitivity_count is not None
>>> assert isinstance(point_sector_sensitivity_count, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q4 (points: 14)
```
>>> assert heatmap_sector_sensitivity_count is not None
>>> assert isinstance(heatmap_sector_sensitivity_count, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: All your variables have the expected name and type


Tests for q6 (points: 10)
```
>>> assert scatter_breaches_over_time is not None
>>> assert isinstance(scatter_breaches_over_time, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: This test only checks that your variables have the correct name and type


Tests for q7 (points: 9)
```
>>> assert bar_method_records_lost is not None
>>> assert isinstance(bar_method_records_lost, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: All your variables have the expected name and type


Tests for q8 (points: 12)
```
>>> assert sbar_method_records_sensitivity is not None
>>> assert isinstance(sbar_method_records_sensitivity, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: All your variables have the expected name and type


Tests for q9 (points: 8)
```
>>> assert scatter_some_sectors is not None
>>> assert isinstance(scatter_some_sectors, alt.Chart)
```
failure_message: Some of your variables have an unexpected name or type
success_message: All your variables have the expected name and type
